<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name>
      <text>_USER_DEFINED_PROTOTYPE_mysql_dockerized_v2.1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>This question template is based on <a href="https://github.com/andre-wojtowicz/coderunner-dba-dataed2026">GRANT Privileges, REVOKE Risk: Safe and Scalable Teaching of Database Administration with Isolated Containers</a> paper by <a href="https://awojt.pl">A. Wójtowicz</a> and M. Prill presented during <a href="https://dataedinitiative.github.io/DataEd26/index.html">DataEd</a> workshop of the <a href="https://edbticdt2026.github.io/">EDBT/ICDT</a> joint conference in 2026.</p>
<h4>MySQL-dockerized</h4>
<p>A teacher working on a new question has to provide a standard question description, an example answer for validation and test cases. The basic setup of a test case requires code in the <strong>Extra template data</strong> section which uses Python functions provided by the template. Verification of a test case is done by textual output comparison. When a database engine reports an error it is displayed to the student.</p>
<p>The following subsections briefly describe how a test case is prepared depending on the question type. This question template starts a fresh isolated Docker container for each test case evaluation and stops it at the end. The container grants full administrative access to the database engine, enabling practice with DCL commands such as <code>GRANT</code> and <code>REVOKE</code> that are not feasible in a shared-server environment. These general guidelines use Python code and can be customized by the teacher accordingly to the course requirements.</p>
<h4>Python Functions</h4>
<p>Two Python functions are available for executing SQL in a test case:</p>
<ul>
<li>
<p><strong><code>invoke_commandline_sql(db_code, database='test_case',
          timeout=6, command_line_tool=COMMAND_LINE_TOOL,
          ignore_warnings=False)</code></strong> - executes SQL by spawning the <code>mysql</code> command-line tool as a subprocess. The <code>command_line_tool</code> parameter can be set to a different binary or argument list when needed. Output is strictly textual. Errors and timeouts terminate the entire script.</p>
</li>
<li>
<p><strong><code>invoke_cursor_sql(db_code, database='test_case',
          ignore_warnings=False, sort_result=False)</code></strong> - executes SQL through a database cursor and formats the result as a pandas DataFrame string. Suitable when structured tabular output is needed.</p>
</li>
</ul>
<p>Both functions accept a <code>database=</code> parameter - when set to a database name it prepends <code>USE &lt;name&gt;;</code> to the SQL; <code>None</code> or <code>False</code> runs at the server default level.</p>
<h4>DQL Question Type</h4>
<p>This type of question uses the <code>SELECT</code> command to extract data from a database. In <strong>Extra template data</strong> section of a test case, firstly, a database is filled with tables, data, etc.; secondly, a student's answer is invoked with output printed on the standard output. Since validation is done in a textual way, the teacher should either a) explicitly put in the question description what columns should the student use to sort results by <code>ORDER BY</code> clause; or b) configure the script to sort rows in a pre-processing step before printing to the standard output. See sample code below:</p>
<pre><code class="language-python">__sql_prepare__ = r'''
-- database init commands e.g. CREATE DATABASE, CREATE TABLE, INSERT etc.
-- ...
'''

invoke_commandline_sql(__sql_prepare__, database=None)
print(invoke_commandline_sql(__student_answer__, database=None))
</code></pre>
<h4>DDL/DML Question Type</h4>
<p>This type of question checks a student's ability to modify a database, e.g. with <code>CREATE</code>, <code>INSERT</code> etc. commands. In <strong>Extra template data</strong> section of a test case, firstly, a database is prepared using the <code>mysql</code> command-line tool (which handles multi-statement initialization including <code>CREATE DATABASE</code>); secondly, a student's answer is invoked in the target database; thirdly, a verification <code>SELECT</code> query is invoked with output printed on standard output. See sample code below:</p>
<pre><code class="language-python">__sql_prepare__ = r'''
-- server-level init e.g. CREATE DATABASE mydb, USE mydb, CREATE TABLE etc.
-- ...
'''

__sql_verify__ = r'''
-- statement checking database structure e.g. SELECT data about
-- columns in a table using INFORMATION_SCHEMA view; test code can be
-- provided here or in 'Test code X' section of a test case
-- ...
'''

invoke_commandline_sql(__sql_prepare__, database=None)
invoke_commandline_sql(__student_answer__, database=None)
print(invoke_commandline_sql(__sql_verify__, database=None))
</code></pre>
<h4>DCL Question Type</h4>
<p>This type of question checks a student's ability to manage access control, e.g. with <code>GRANT</code>, <code>REVOKE</code> etc. commands. In <strong>Extra template data</strong> section of a test case, firstly, the necessary users, databases and objects are created; secondly, a student's answer is invoked; thirdly, a verification query inspects the privilege state and prints it on standard output. See sample code below:</p>
<pre><code class="language-python">__sql_prepare__ = r'''
-- init commands e.g. CREATE DATABASE mydb, CREATE USER 'myuser'@'%',
-- USE mydb, CREATE TABLE etc.
-- ...
'''

__sql_verify__ = r'''
-- statement checking privilege state e.g. SHOW GRANTS FOR 'myuser'@'%'
-- or SELECT from information_schema.USER_PRIVILEGES etc.
-- ...
'''

invoke_commandline_sql(__sql_prepare__, database=None)
invoke_commandline_sql(__student_answer__, database=None)
print(invoke_commandline_sql(__sql_verify__, database=None))
</code></pre>
<h4>Other types of assignments</h4>
<p>Above we described DQL, DML, DDL and DCL queries but our solution is able to cover problems where a student has to write e.g. SQL script, stored procedures, functions, backup and restore operations etc., as long as created objects persist within the duration of the Docker container for that test case. In the above examples we used <code>invoke_commandline_sql(...)</code> which uses by default the <code>mysql</code> command-line tool and produces strictly textual output. For questions where programmatic investigation of structured tabular output is needed, <code>invoke_cursor_sql(...)</code> can be used instead as it processes data through a database cursor.</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>mysql_dockerized_v2.1</coderunnertype>
    <prototypetype>2</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace>1</useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[# v2026-03-18
import sys
import site
import os
venv_path = "/opt/venv/mysql-dockerized"
sys.path.insert(0, os.path.join(venv_path, "lib", "python3.11", "site-packages"))
site.addsitedir(os.path.join(venv_path, "lib", "python3.11", "site-packages"))
os.environ["VIRTUAL_ENV"] = venv_path
os.environ["PATH"] = f"{venv_path}/bin:" + os.environ["PATH"]

import time
import getpass
import subprocess
import pandas
import docker
import mysql.connector
import _mysql_connector # for catching errors when fetching data

# The server port to use. When running without Jobe this is just '140', but
# when in Jobe it'll become a combination of the current Jobe user.
# For jobe00 it'll be 14000, for jobe01 14001 and so on.
#
# When there are multiple Docker question types this must be unique, and
# the combined number must be lower than max port number - 65535.
PORT = int('139'+''.join(filter(str.isdigit, getpass.getuser())))
DOCKER_IMAGE = 'jobe-mysql:8.0.45-1'
DOCKER_NAME = 'jobe-mysql_' + str(PORT)
DOCKER_PORT = {'3306/tcp': PORT}
DOCKER_TMPFS = {'/var': 'size=512M'}
CFG = {
    'user': 'root',
    'password': 'jobe_runner',
    'host': 'localhost',
    'port': PORT,
    'autocommit': True,
    'buffered': True
    #'get_warnings': True # disabled due to a bug (https://forums.mysql.com/read.php?50,677572,677580#msg-677580)
}
COMMAND_LINE_TOOL = [
    '/usr/bin/mysql',
    '--host='+CFG['host'],
    '--port='+str(CFG['port']),
    '--user='+CFG['user'],
    '--connect_timeout=5',
    '--protocol=TCP',
    '--table'
    ]

START_TIME = time.time()
def DEBUG_TIME(string):
    """ Outputs time since last invocation with a description, used for debugging. """
    global START_TIME
    #print("- %s - %s" % (string, (time.time() - START_TIME)))
    START_TIME = time.time()

def create_container():
    """ Checks if there is a created docker instance with the right image version. """
    docker_client = docker.from_env()
    try:
        curr = docker_client.containers.get(DOCKER_NAME)
        if curr.image.tags[0] != DOCKER_IMAGE:
            curr.remove(force=True)
            DEBUG_TIME('recreate')
            raise docker.errors.NotFound("wrong version")
        return curr
    except docker.errors.NotFound:
        # if it doesn't then create a new one
        DEBUG_TIME('create')
        return docker_client.containers.create(
            image=DOCKER_IMAGE,
            name=DOCKER_NAME,
            read_only=True,
            ports=DOCKER_PORT,
            tmpfs=DOCKER_TMPFS
            )

def wait_for_connection(timeout=7.0):
    """ Attempts to connect to the server with a backoff time until timeout is reached. """
    while True:
        try:
            return mysql.connector.connect(**CFG)
        except BaseException as ex:
            time.sleep(0.1)
            if time.time() - START_TIME >= timeout:
                raise TimeoutError('Waited too long for server to start.') from ex

def cleanup():
    """ Cleans up resources. """
    cursor.close()
    connection.close()
    container.stop(timeout=0)
    DEBUG_TIME("stop")

def fail_with_error(error_string):
    print(error_string, file=sys.stderr)
    cleanup()
    sys.exit(1)

def invoke_commandline_sql(db_code,
                           database='test_case',
                           timeout=6,
                           command_line_tool=COMMAND_LINE_TOOL,
                           ignore_warnings=False):
    """ Invokes given SQL via a command line tool.
    Errors and timeout will exit the entire script.

    Keyword arguments:
    db_code -- SQL to invoke
    database -- database to execute the query in, can be False to disable the feature
    timeout -- timeout which MUST BE LESS THAN DEFAULT FOR QUESTION TYPE
    ignore_warnings -- whether to ignore warnings, by default False
    """
    if database:
        db_code = "USE %s;\n" % database + db_code
    if not ignore_warnings:
        command_line_tool.append('--show-warnings')
    try:
        # to stop mysql from putting a warning about command line password
        # into stderr we create a copy of the host environment and put the password there
        mysqlenv = os.environ.copy()
        mysqlenv['MYSQL_PWD'] = CFG['password']
        db_outcome = subprocess.run(
            command_line_tool,
            input=db_code,
            stdout=subprocess.PIPE,
            stderr=subprocess.PIPE,
            timeout=timeout,
            universal_newlines=True,
            check=False,
            env=mysqlenv
        )
    except subprocess.TimeoutExpired as e:
        fail_with_error(f'''SQL invocation timed out: {e}''')
    if db_outcome.returncode != 0:
        fail_with_error(f'''SQL invocation failed
Process exit code: {db_outcome.returncode}
Server response:
{db_outcome.stdout}
{db_outcome.stderr}
''')
    # return code == 0 and stderr present means there's warnings
    # from the mysql command, but warnings from the mysql server
    # are in standard output if --show-warnings is specified
    if db_outcome.stderr and not ignore_warnings:
        return f'''{db_outcome.stdout}\n{db_outcome.stderr}'''
    return db_outcome.stdout

def invoke_cursor_sql(db_code, database='test_case', ignore_warnings=False, sort_result=False):
    """ Invokes given SQL via the cursor.
    Warnings will be appended to result.
    Errors will exit the entire script.

    Keyword arguments:
    db_code -- SQL to invoke
    database -- database to execute the query in, can be False to disable the feature
    ignore_warnings -- whether to ignore warnings, by default False
    sort_result -- output sorted (by sorted list of all columns, ascending) results, by default False
    """
    try:
        if database:
            db_code = "USE %s;\n" % database + db_code
        result = cursor.execute(db_code, multi=True)
    except mysql.connector.Error as e:
        fail_with_error(e)

    warnings = ''
    try:
        for r in result:
            for w in cursor.fetchwarnings() or []:
                warnings = f'''{warnings}{" ".join(map(str, w))}\n'''
            if r.with_rows:
                # handle selective encoding of strings into bytes arrays
                rows = []
                for list in r:
                    rows.append([el.decode('utf-8') if type(el) is bytes else el for el in list])
                cols = [x[0] for x in r.description]
                df = pandas.DataFrame(rows, range(1, len(rows)+1), cols)
                if sort_result:
                    df.sort_values(by=list(df.columns), inplace=True)
                    answer_result = df[sorted(cols)].fillna(value="NULL").to_string(index=False)
                else:
                    answer_result = df.fillna(value="NULL").to_string()
                break
            answer_result = "(no data)"
    # mysql cursor API doesn't specify _mysql_connector.MySQLInterfaceError 
    # as expected here, but it can be raised anyway
    except (mysql.connector.Error, _mysql_connector.MySQLInterfaceError) as e:
        fail_with_error(e)

    if warnings and not ignore_warnings:
        answer_result = '{0}\n{1}'.format(answer_result, warnings.strip())

    return answer_result

__student_answer__ = """{{ STUDENT_ANSWER }}"""

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
container = create_container()
DEBUG_TIME('client')

# if it's running stop it
# since everything runs in tmpfs with read-only root this resets the state
if container.status == 'running':
    container.stop(timeout=0)
container.start()
DEBUG_TIME('start')

# wait for connection to be ready
connection = wait_for_connection()
DEBUG_TIME('connection')

cursor = connection.cursor()

__testcode__ = r"""{{ TEST.testcode }}"""
{{ TEST.extra }}

cleanup()

{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}
]]></template>
    <iscombinatortemplate>1</iscombinatortemplate>
    <allowmultiplestdins>0</allowmultiplestdins>
    <answer></answer>
    <validateonsave>0</validateonsave>
    <testsplitterre><![CDATA[|#<ab@17943918#@>#\n|ms]]></testsplitterre>
    <language>python3</language>
    <acelang>MySQL</acelang>
    <sandbox></sandbox>
    <grader>EqualityGrader</grader>
    <cputimelimitsecs>30</cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin>ace</uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
    </testcases>
  </question>

</quiz>
