<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name>
      <text>01 - list all databases</text>
    </name>
    <questiontext format="markdown">
      <text>Write a psql meta-command that lists all databases on the server.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>3</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>\l</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>                                              List of databases
    Name     | Owner | Encoding |  Collate   |   Ctype    | ICU Locale | Locale Provider | Access privileges
-------------+-------+----------+------------+------------+------------+-----------------+-------------------
 jobe_runner | root  | UTF8     | en_US.utf8 | en_US.utf8 |            | libc            |
 postgres    | root  | UTF8     | en_US.utf8 | en_US.utf8 |            | libc            |
 template0   | root  | UTF8     | en_US.utf8 | en_US.utf8 |            | libc            | =c/root          +
             |       |          |            |            |            |                 | root=CTc/root
 template1   | root  | UTF8     | en_US.utf8 | en_US.utf8 |            | libc            | =c/root          +
             |       |          |            |            |            |                 | root=CTc/root
(4 rows)
</text>
      </expected>
      <extra>
                <text>print(invoke_commandline_sql(__student_answer__))
</text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>02 - create roles and users</text>
    </name>
    <questiontext format="markdown">
      <text>Create roles named `reader` and `writer` (no login). Also create users `alice` and `bob` (with login capability).</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>6</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>CREATE USER alice;
CREATE USER bob;
CREATE ROLE reader;
CREATE ROLE writer;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>                                   List of roles
 Role name |                         Attributes                         | Member of
-----------+------------------------------------------------------------+-----------
 alice     |                                                            | {}
 bob       |                                                            | {}
 reader    | Cannot login                                               | {}
 root      | Superuser, Create role, Create DB, Replication, Bypass RLS | {}
 writer    | Cannot login                                               | {}</text>
      </expected>
      <extra>
                <text>
__sql_verify__ = '''
\du
'''
invoke_commandline_sql(__student_answer__)
print(invoke_commandline_sql(__sql_verify__))
</text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>03 - grant CONNECT on database</text>
    </name>
    <questiontext format="markdown">
      <text>Grant the `CONNECT` privilege on database `appdb` to roles `reader` and `writer`.
Also revoke `CONNECT` from `PUBLIC` so that only explicitly granted roles can connect.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>6</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>GRANT CONNECT ON DATABASE appdb TO reader;
GRANT CONNECT ON DATABASE appdb TO writer;
REVOKE CONNECT ON DATABASE appdb FROM PUBLIC;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>                                           List of databases
 Name  | Owner | Encoding |  Collate   |   Ctype    | ICU Locale | Locale Provider | Access privileges
-------+-------+----------+------------+------------+------------+-----------------+-------------------
 appdb | root  | UTF8     | en_US.utf8 | en_US.utf8 |            | libc            | =T/root          +
       |       |          |            |            |            |                 | root=CTc/root    +
       |       |          |            |            |            |                 | reader=c/root    +
       |       |          |            |            |            |                 | writer=c/root
(1 row)
</text>
      </expected>
      <extra>
                <text>__sql_prepare__ = '''
CREATE ROLE reader;
CREATE ROLE writer;
CREATE DATABASE appdb;
'''
invoke_commandline_sql(__sql_prepare__)
invoke_commandline_sql(__student_answer__)

__sql_verify__ = '''
\l appdb
'''
print(invoke_commandline_sql(__sql_verify__))
</text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>04 - grant table privileges</text>
    </name>
    <questiontext format="markdown">
      <text>Grant role `reader` the `SELECT` privilege on both tables `products` and `orders` in database `appdb`.
Grant role `writer` the `SELECT`, `INSERT`, `UPDATE` and `DELETE` privileges on both tables.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>6</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>GRANT SELECT ON TABLE products, orders TO reader;
GRANT SELECT, INSERT, UPDATE, DELETE ON TABLE products, orders TO writer;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>                              Access privileges
 Schema |   Name   | Type  | Access privileges | Column privileges | Policies
--------+----------+-------+-------------------+-------------------+----------
 public | orders   | table | root=arwdDxt/root+|                   |
        |          |       | reader=r/root    +|                   |
        |          |       | writer=arwd/root  |                   |
 public | products | table | root=arwdDxt/root+|                   |
        |          |       | reader=r/root    +|                   |
        |          |       | writer=arwd/root  |                   |
(2 rows)
</text>
      </expected>
      <extra>
                <text><![CDATA[__sql_prepare__ = '''
CREATE ROLE reader;
CREATE ROLE writer;
CREATE DATABASE appdb;
\c appdb
CREATE TABLE products (id INT PRIMARY KEY, name TEXT, price FLOAT);
CREATE TABLE orders   (id INT PRIMARY KEY, product_id INT, qty INT);
INSERT INTO products VALUES (1, 'Widget', 9.99), (2, 'Gadget', 24.99);
INSERT INTO orders   VALUES (1, 1, 3), (2, 2, 1);
GRANT CONNECT ON DATABASE appdb TO reader;
GRANT CONNECT ON DATABASE appdb TO writer;
REVOKE CONNECT ON DATABASE appdb FROM PUBLIC;
'''
invoke_commandline_sql(__sql_prepare__)

CFG['dbname'] = 'appdb'
COMMAND_LINE_TOOL = [
    '/usr/bin/psql',
    '--set=ON_ERROR_STOP=1',
    'postgresql://' \
    + CFG['user']+':'+CFG['password']+'@' \
    + CFG['host']+':'+str(CFG['port'])+'/'+CFG['dbname'] \
    + '?connect_timeout=2&sslmode='+CFG['sslmode']
]
invoke_commandline_sql(__student_answer__, command_line_tool=COMMAND_LINE_TOOL)

__sql_verify__ = '''
\dp
'''
print(invoke_commandline_sql(__sql_verify__, command_line_tool=COMMAND_LINE_TOOL))
]]></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>05 - revoke CONNECT from user</text>
    </name>
    <questiontext format="markdown">
      <text>Revoke the `CONNECT` privilege on database `appdb` from user `alice`.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>3</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>REVOKE CONNECT ON DATABASE appdb FROM alice;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>                                           List of databases
 Name  | Owner | Encoding |  Collate   |   Ctype    | ICU Locale | Locale Provider | Access privileges
-------+-------+----------+------------+------------+------------+-----------------+-------------------
 appdb | root  | UTF8     | en_US.utf8 | en_US.utf8 |            | libc            | =T/root          +
       |       |          |            |            |            |                 | root=CTc/root    +
       |       |          |            |            |            |                 | reader=c/root
(1 row)
</text>
      </expected>
      <extra>
                <text>__sql_prepare__ = '''
CREATE USER alice;
CREATE ROLE reader;
CREATE DATABASE appdb;
GRANT CONNECT ON DATABASE appdb TO alice;
GRANT CONNECT ON DATABASE appdb TO reader;
REVOKE CONNECT ON DATABASE appdb FROM PUBLIC;
'''
invoke_commandline_sql(__sql_prepare__)
invoke_commandline_sql(__student_answer__)

__sql_verify__ = '''
\l appdb
'''
print(invoke_commandline_sql(__sql_verify__))
</text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>06 - create schema and change owner</text>
    </name>
    <questiontext format="markdown">
      <text>In database `appdb`, create a schema named `sales`. Then change its owner to user `alice`.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>4</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>CREATE SCHEMA sales;
ALTER SCHEMA sales OWNER TO alice;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>List of schemas
  Name  | Owner
--------+-------
 public | root
 sales  | alice
(2 rows)
</text>
      </expected>
      <extra>
                <text><![CDATA[__sql_prepare__ = '''
CREATE USER alice;
CREATE DATABASE appdb;
\c appdb
'''
invoke_commandline_sql(__sql_prepare__)

CFG['dbname'] = 'appdb'
COMMAND_LINE_TOOL = [
    '/usr/bin/psql',
    '--set=ON_ERROR_STOP=1',
    'postgresql://' \
    + CFG['user']+':'+CFG['password']+'@' \
    + CFG['host']+':'+str(CFG['port'])+'/'+CFG['dbname'] \
    + '?connect_timeout=2&sslmode='+CFG['sslmode']
]
invoke_commandline_sql(__student_answer__, command_line_tool=COMMAND_LINE_TOOL)

__sql_verify__ = '''
\dn
'''
print(invoke_commandline_sql(__sql_verify__, command_line_tool=COMMAND_LINE_TOOL))
]]></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>07 - grant schema privileges</text>
    </name>
    <questiontext format="markdown">
      <text>Grant role `reader` the `USAGE` privilege on schema `sales` and `SELECT` on all existing tables in that schema.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>4</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>GRANT USAGE ON SCHEMA sales TO reader;
GRANT SELECT ON ALL TABLES IN SCHEMA sales TO reader;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>                              Access privileges
 Schema |  Name   | Type  | Access privileges | Column privileges | Policies
--------+---------+-------+-------------------+-------------------+----------
 sales  | regions | table | root=arwdDxt/root+|                   |
        |         |       | reader=r/root     |                   |
(1 row)
</text>
      </expected>
      <extra>
                <text><![CDATA[__sql_prepare__ = '''
CREATE ROLE reader;
CREATE DATABASE appdb;
\c appdb
CREATE SCHEMA sales;
CREATE TABLE sales.regions (id INT PRIMARY KEY, name TEXT);
INSERT INTO sales.regions VALUES (1, 'North'), (2, 'South');
'''
invoke_commandline_sql(__sql_prepare__)

CFG['dbname'] = 'appdb'
COMMAND_LINE_TOOL = [
    '/usr/bin/psql',
    '--set=ON_ERROR_STOP=1',
    'postgresql://' \
    + CFG['user']+':'+CFG['password']+'@' \
    + CFG['host']+':'+str(CFG['port'])+'/'+CFG['dbname'] \
    + '?connect_timeout=2&sslmode='+CFG['sslmode']
]
invoke_commandline_sql(__student_answer__, command_line_tool=COMMAND_LINE_TOOL)

__sql_verify__ = '''
\dp sales.*
'''
print(invoke_commandline_sql(__sql_verify__, command_line_tool=COMMAND_LINE_TOOL))
]]></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>08 - create tablespace</text>
    </name>
    <questiontext format="markdown">
      <text>Create a tablespace named `fast` located at `/var/abd/tablespace2`.
The directory has already been prepared.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>3</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>CREATE TABLESPACE fast LOCATION '/var/abd/tablespace2';</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>            List of tablespaces
    Name    | Owner |       Location
------------+-------+----------------------
 big        | root  | /var/abd/tablespace1
 fast       | root  | /var/abd/tablespace2
 pg_default | root  |
 pg_global  | root  |
(4 rows)</text>
      </expected>
      <extra>
                <text>
__sql_prepare__ = '''
CREATE TABLESPACE big LOCATION '/var/abd/tablespace1';
'''
invoke_commandline_sql(__sql_prepare__)

invoke_commandline_sql(__student_answer__)

__sql_verify__ = '''
\db
'''
print(invoke_commandline_sql(__sql_verify__))
</text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>09 - create database in tablespace and grant CREATE</text>
    </name>
    <questiontext format="markdown">
      <text>Create a database named `telemetry` in tablespace `fast`. The owner should be role `dba`.
Then grant the `CREATE` privilege on tablespace `big` to role `dba`.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>4</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>CREATE DATABASE telemetry OWNER dba TABLESPACE fast;
GRANT CREATE ON TABLESPACE big TO dba;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>                                             List of databases
   Name    | Owner | Encoding |  Collate   |   Ctype    | ICU Locale | Locale Provider | Access privileges
-----------+-------+----------+------------+------------+------------+-----------------+-------------------
 telemetry | dba   | UTF8     | en_US.utf8 | en_US.utf8 |            | libc            |
(1 row)

                                    List of tablespaces
 Name | Owner |       Location       | Access privileges | Options |  Size   | Description
------+-------+----------------------+-------------------+---------+---------+-------------
 big  | root  | /var/abd/tablespace1 | root=C/root      +|         | 0 bytes |
      |       |                      | dba=C/root        |         |         |
(1 row)
</text>
      </expected>
      <extra>
                <text>__sql_prepare__ = '''
CREATE TABLESPACE big LOCATION '/var/abd/tablespace1';
CREATE TABLESPACE fast LOCATION '/var/abd/tablespace2';
CREATE USER dba WITH PASSWORD 'dba';
'''
invoke_commandline_sql(__sql_prepare__)
invoke_commandline_sql(__student_answer__)

__sql_verify__ = '''
\l telemetry
\db+ big
'''
print(invoke_commandline_sql(__sql_verify__))
</text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name>
      <text>10 - move table to tablespace</text>
    </name>
    <questiontext format="markdown">
      <text>You are connected to database `telemetry` as role `dba`.

Move table `archive` to tablespace `fast`.</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>pgsql_dockerized_v2.1</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>3</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>ALTER TABLE archive SET TABLESPACE fast;</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text><![CDATA[               Table "public.archive"
 Column  |  Type   | Collation | Nullable | Default
---------+---------+-----------+----------+---------
 id      | integer |           |          |
 user_id | integer |           |          |
 url     | text    |           |          |
]]></text>
      </expected>
      <extra>
                <text><![CDATA[__sql_prepare__ = '''
CREATE TABLESPACE big LOCATION '/var/abd/tablespace1';
CREATE TABLESPACE fast LOCATION '/var/abd/tablespace2';
CREATE USER dba WITH PASSWORD 'dba';
CREATE DATABASE telemetry OWNER dba TABLESPACE fast;
GRANT CREATE ON TABLESPACE big TO dba;
GRANT CREATE ON TABLESPACE fast TO dba;
'''
invoke_commandline_sql(__sql_prepare__)

CFG['user'] = 'dba'
CFG['password'] = 'dba'
CFG['dbname'] = 'telemetry'
COMMAND_LINE_TOOL = [
    '/usr/bin/psql',
    '--set=ON_ERROR_STOP=1',
    'postgresql://' \
    + CFG['user']+':'+CFG['password']+'@' \
    + CFG['host']+':'+str(CFG['port'])+'/'+CFG['dbname'] \
    + '?connect_timeout=2&sslmode='+CFG['sslmode']
]
__sql_prepare__ = '''
CREATE TABLE archive (id INT, user_id INT, url TEXT) TABLESPACE big;
'''
invoke_commandline_sql(__sql_prepare__, command_line_tool=COMMAND_LINE_TOOL)
invoke_commandline_sql(__student_answer__, command_line_tool=COMMAND_LINE_TOOL)

__sql_verify__ = '''
\d archive
'''
print(invoke_commandline_sql(__sql_verify__, command_line_tool=COMMAND_LINE_TOOL))
]]></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>
